# PS03
Name: YOUR NAME HERE 
Collaborators:
 - Collaborator 1
 - Collaborator 2

DS256, Gettysburg College  
Prof Eatai Roth  
**Due: Wed, Feb 18, 2026 at 1:10p**

By submitting this work, you attest that: 
 - it reflects your own effort and understanding of the material
 - you did not ask for help from students outside the course, other than PLAs
 - you did not use a generative AI to produce any code used below.

 

The questions are in no particular order in terms of difficulty, so if you feel stuck, skip the question and come back to it.

**1.** Write a function ```factorize()``` that takes as input an integer and returns as output a list of factors of that number. If the number is prime, also print "That's prime."

*Useful tool:* The function ```range(start, stop, step)``` creates a sequence of numbers as a list.

```python
def factorize(num):
    return

```

```python
# Is it doing what you expect it to do.
print(factorize(1)) # should be [1]
print(factorize(57)) # should be [1, 3, 19, 57]
print(factorize(59)) # should print "That's prime." and return [1, 59]
```

Output:
```
None
None
None
```

**2.**

Write the function ```same_set``` that takes two lists and checks if they have all the same elements (and the same number of any repeated elements) and returns ```True``` if the contents match. For example, the lists:
 -  ```A = [2, 'legit', 2, 'quit']```
 -  ```B = [2, 2, 'quit', 'legit']```
 -  ```C = ['quit', 2, 'legit']```
 -  ```D = ['too', 'legit', 'to', 'quit']```
 
only ```same_set(A,B)``` would be ```True```.

```python
'''your code here'''
```

Output:
```
'your code here'
```

**3.**

Write the function ```check_cyclic_equality(A, B)``` that takes two lists and checks if there is a rotation that makes them same. If the lists are cyclically equal, return the smallest number the second list has to be cycled to equal the first list. If not, return ```False```.

 - A single rotation/cycling is defined as taking the zero-th element of the list and moving it to the end. For example, [‘dog’, 7, ‘chicken’, ‘rabbit’, 9] cycled by 2 would be [‘chicken’, ‘rabbit’, 9, ‘dog’, 7].
 - Two lists CANNOT be cyclically equal if they don't contain all of the same elements, so maybe check that first with the ```same_set``` function you wrote above.

```python
def check_cyclic_equality(A, B):
    return

```

```python
# Lists are cyclically equal with 2 rotations
assert check_cyclic_equality([1, 2, 3], [3, 1, 2]) == 1

# Lists are not cyclically equal
assert check_cyclic_equality([1, 2, 3], [4, 5, 6]) == False
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[10], line 2
      1 # Lists are cyclically equal with 2 rotations
----> 2 assert check_cyclic_equality([1, 2, 3], [3, 1, 2]) == 1
      4 # Lists are not cyclically equal
      5 assert check_cyclic_equality([1, 2, 3], [4, 5, 6]) == False

AssertionError:
```

**4.** On the game show Wheel of Fortune, contestants are presented a text puzzle where all the letters are blank. They guess letters to fill in the blanks until they can decode the puzzle. In the final round, the winning contestant is presented a puzzle and granted the letters R, S, T, L, N, E for free (every other letter appears as a blank). 

Write the function ```wheel_final``` that takes in a string of text and returns a string with every letter other than R, S, T, L, N, E replaced with a blank. Spaces and punctuation should also be left alone. All letters should be captilized (using ```.upper()```).

For example: ```wheel_final('Gettysburg College`)``` should return ```"_ETT_S__R_ __LLE_E"```.

*Hints*: 
- You can consider a string a list of characters and loop over the characters.

```python
hi = 'Hello!'
for letter in hi:
    print(letter)
    
'Gettysburg College'.lower()
```

**5.** 

Write a function, ```find_filetype()``` which will take as input a list of filenames and a file extension starting with a '.' (e.g. '.pdf', '.xlsx', '.ipynb') and return the sublist of only those files with the specified filetype.

```python
def find_filetype(files):
    return
    
```

```python
file_list = ['hi.pdf', 'hello.doc', 'ps01.ipynb', 'goodbye.pdf', 'test.ipynb']

find_filetype(file_list, '.ipynb')
```

**6.** Write the function ```remove_duplicates()``` that takes as an input a list or string and removes any duplicates, returning a list or string of only unique values in the order they appear in the original list.

*Hint*: As you build your list of unique values, you'll be checking whether new values are already in that list.

```python
def remove_duplicates(entry):
    return
```

```python
assert remove_duplicates([9, 1, 1, 'Emergency']) == [9, 1, 'Emergency']
assert remove_duplicates('Hello there!') == 'Helo thr!'
```

**7. (Optional challenge)** Write the function ```convert_to_numeric()``` that takes as input a list of prices. Some prices are listed as numbers, others as strings. Strings may appear with or without a decimal (but guaranteed that only one decimal per string) and possibly with a '$' or 'USD'). Return a list of prices as floats. If there is a non-numeric entry (e.g. 'Out of Stock'), skip it.

*Hint*: 
- In class we checked the type of a variable as ```type(var)==int```. The function ```isinstance(var, [types])``` can check whether a variable is of several types.

```isinstance(12.43, [int, float])``` --> True

```python
'''your code here'''

```

```python
# Some tests
assert convert_to_numeric([12, '15.50', '$20', '30 USD', 'Out of Stock']) == [12.0, 15.5, 20.0, 30.0]
assert convert_to_numeric(['5', 7.25, '$8.00', 'Not available']) == [5.0, 7.25, 8.0]

```

Tests for q1 (points: None)
```
>>> assert factorize(57) == [1, 3, 19, 57]
>>> assert factorize(59) == [1, 59]
>>> assert factorize(1) == [1], "Your function doesn't handle the 1 case correctly."
>>> assert factorize(-7) is None, 'Remember to check the number is positive.'
That's prime.
Please use a positive integer.
```


Tests for q2 (points: None)
```
>>> assert check_cyclic_equality([1, 2, 3], [1, 2, 3]) == 0
>>> assert check_cyclic_equality([1, 2, 3], [2, 3, 1]) == 2
>>> assert check_cyclic_equality([1, 2, 3], [1, 2]) == False
>>> assert check_cyclic_equality([], []) == 0
>>> assert check_cyclic_equality(['a', 'b', 'a'], ['b', 'a', 'a']) == 2
>>> assert check_cyclic_equality(['dog', 'cat', 'bird'], ['cat', 'bird', 'dog']) == 2
```


Tests for q3 (points: None)
```
>>> assert extract_vowels('aeiou') == [1, 1, 1, 1, 1, 5]
>>> assert extract_vowels('AEIOU') == [1, 1, 1, 1, 1, 5]
>>> assert extract_vowels('AaBbCcDdEeFfGgHhIiJjKkLlMmNnOoPpQqRrSsTtUuVvWwXxYyZz') == [2, 2, 2, 2, 2, 10]
>>> assert extract_vowels('bcdfg') == [0, 0, 0, 0, 0, 0]
>>> assert extract_vowels('Gettysburg College') == [1, 2, 1, 2, 1, 7]
>>> assert extract_vowels('') == [0, 0, 0, 0, 0, 0]
>>> assert extract_vowels('123!@#aei') == [1, 1, 1, 0, 0, 3]
```


Tests for q4 (points: None)
```
>>> file_list = ['hi.pdf', 'hello.doc', 'ps01.ipynb', 'goodbye.pdf', 'test.ipynb']
>>> assert find_filetype(file_list, '.pdf') == ['hi.pdf', 'goodbye.pdf']
>>> assert find_filetype(file_list, '.ipynb') == ['ps01.ipynb', 'test.ipynb']
>>> assert find_filetype(file_list, '.doc') == ['hello.doc']
>>> assert find_filetype(file_list, '.xlsx') == []
>>> assert find_filetype([], '.pdf') == []
```


Tests for q5 (points: None)
```
>>> assert remove_duplicates([1, 2, 2, 3, 1, 4]) == [1, 2, 3, 4]
>>> assert remove_duplicates([5, 6, 7]) == [5, 6, 7]
>>> assert remove_duplicates([9, 9, 9, 9]) == [9]
>>> assert remove_duplicates([]) == []
>>> assert remove_duplicates(['a', 'b', 'a', 'c', 'b']) == ['a', 'b', 'c']
>>> assert remove_duplicates('A man, a plan, a canal, Panama.') == ['A', ' ', 'm', 'a', 'n', ',', 'p', 'l', 'c', 'P', '.']
```


Tests for q6 (points: None)
```
>>> assert convert_to_numeric([10, '20', 30.5, '$40', '50 USD']) == [10.0, 20.0, 30.5, 40.0, 50.0]
>>> assert convert_to_numeric([1, 2.5, 3]) == [1.0, 2.5, 3.0]
>>> assert convert_to_numeric(['$5', '6 USD', '$7.50']) == [5.0, 6.0, 7.5]
>>> assert convert_to_numeric(['Out of Stock', 'N/A', 'Unavailable']) == []
>>> assert convert_to_numeric([100, '200', 'Not available', '$300', '400 USD', 'error']) == [100.0, 200.0, 300.0, 400.0]
>>> assert convert_to_numeric([]) == []
>>> assert convert_to_numeric(['1.99', '$2.50', '3.00 USD']) == [1.99, 2.5, 3.0]
```
