# PS01
Name: <YOUR NAME HERE>

DS256, Gettysburg College  
Prof Eatai Roth  
**Due: Wed, Feb 4, 2026 at 1:10p**

## Instructions
 - Write your code in the cells labeled '''your code here'''.
 - The questions are in no particular order in terms of difficulty, so if you feel stuck, skip the question and come back to it.

**1.** Write a function named ```hi_there``` that prints the phrase "Hi there!". This function takes no inputs and has no outputs.

```python
'''your code here'''
```

Output:
```
'your code here'
```

```python
# If you see the friendly message below this cell, you've done it correctly.
hi_there()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[2], line 2
      1 # If you see the friendly message below this cell, you've done it correctly.
----> 2 hi_there()

NameError: name 'hi_there' is not defined
```

**2.** Write a function called ```last_name_first```. The function should take two inputs named 'first_name', and 'last_name' and **return** a string that is 'last_name, first_name'. If either first or last name is missing, substitute a question mark (?) for the missing name.

```python
'''your code here'''
```

Next, I provide a few tests to check if your code is doing what it should. In the (very near) future, you will be writing these tests yourself.

```python
print(last_name_first('Eatai', 'Roth')) # should print 'Roth, Eatai'
print(last_name_first(None, 'Roth')) # should print 'Roth, ?'
print(last_name_first()) # should print '?, ?'

```

**3.** Write the function ```echo``` that takes as inputs 'shout' (a string) and 'repeats' (an integer). The function should **return** a string that repeats 'shout' the specified number of times. Between each repeat, there should be as many spaces as there are repeats; that is the more repeats, the more spaced out the words will be. There should be no trailing spaces; the output should end with the last instance of the shout word. If repeats is 0 or negative, the function should still return the original shout word.

The function will not print the output, just return the string value.

```python
'''your code here'''
```

```python
print(echo('bye', -1))
```

Another way to test your code is with ```assert``` statements. If the statement after ```assert``` is ```True```, nothing happens and the code progresses; if the statement is ```False```, you will get an "assertion error".

Assertions serve as an alarm that something has gone wrong. Developers use assertions in code in order to protect a desired behavior from being destroyed by future changes to the code.

```python
assert echo('hello', 4) == 'hello    hello    hello    hello'
assert echo('bye', 0) == 'bye'
assert echo('bye', -2) == 'bye'

print('Yay! If this message is displayed, it means all of the above tests have passed!')
```

**4.** Write the function ```banana``` that takes as an argument 'nanum' and returns the word ```banana``` with the number of na's specified. If the function is called with no input, return the word 'banana' spelled as normal.

For example, ```banana(1)``` returns 'bana'.

```python
'''your code here'''
```

```python
# Tests

assert banana(0) == 'ba'
assert banana(1) == 'bana'
assert banana(-2) == 'ba'
assert banana(3) == 'bananana'
assert banana(5) == 'bananananana'
```

**From here on out, you will write your own tests!**

**5.** Write the function ```four_letter_word``` that takes as an input 'word'. If the word is four letters long, return True, otherwise False. The function len() returns the length of a string (e.g. len('hello') --> 5).

```python
'''your code here'''
```

**6.** Write the function ```pos_neg``` that takes as an argument a number and returns the string 'positive' if the number is positive, 'negative' if the number is negative, and 'zero' if the number is exactly zero. 

*Hint:* When used in an equation along with numbers or strings, booleans are interpreted as the integers 1 (True) and 0 (False).

For those who have programmed before, DO NOT USE CONDITIONAL ('if' or 'match-case') STATEMENTS.

```python
'''your code here'''

```

**7.** Write the function ```greater``` that takes as input two numbers and returns the greater of the two. If the two numbers are equal, just return that number.

```python
'''your code here'''
```

**8.** Write the function ```check_square``` that takes two numbers as inputs. If the second number is the first number squared, return True, otherwise False.

```python
'''your code here'''
```

**9.** Write the function ```area_perim``` that takes as input the height and width of a rectangle and returns the area and perimeter of the rectangle (in that order).

```python
'''your code here'''
```

**10.** Write the function ```make_negative``` that takes a number as an input and returns the number as negative. If the number is already negative, just leave it as is.

Again, for those who have learned about conditional statements, try this question without.

```python
'''your code here'''
```

Tests for q10 (points: None)
```
>>> assert make_negative(5) == -5
>>> assert make_negative(-3) == -3
>>> assert make_negative(0) == 0
>>> assert make_negative(100) == -100
>>> assert make_negative(-0.5) == -0.5
```


Tests for q2 (points: None)
```
>>> assert last_name_first('Eatai', 'Roth') == 'Roth, Eatai'
>>> assert last_name_first('Eatai') == '?, Eatai'
>>> assert last_name_first(last_name='Roth') == 'Roth, ?'
>>> assert last_name_first() == '?, ?'
>>> assert last_name_first('', '') == ', '
```


Tests for q3 (points: None)
```
>>> assert echo('hello', 4) == 'hello    hello    hello    hello'
>>> assert echo('bye', 0) == 'bye'
>>> assert echo('bye', -2) == 'bye'
```


Tests for q4 (points: None)
```
>>> assert banana(0) == 'ba'
>>> assert banana(1) == 'bana'
>>> assert banana(-2) == 'ba'
>>> assert banana(3) == 'bananana'
>>> assert banana(5) == 'bananananana'
```


Tests for q5 (points: None)
```
>>> assert four_letter_word('test') == True
>>> assert four_letter_word('word') == True
>>> assert four_letter_word('hi') == False
>>> assert four_letter_word('hello') == False
>>> assert four_letter_word('') == False
```


Tests for q6 (points: None)
```
>>> assert pos_neg(5) == 'positive'
>>> assert pos_neg(-3) == 'negative'
>>> assert pos_neg(0) == 'zero'
>>> assert pos_neg(100) == 'positive'
>>> assert pos_neg(-0.1) == 'negative'
```


Tests for q7 (points: None)
```
>>> assert greater(5, 3) == 5
>>> assert greater(2, 8) == 8
>>> assert greater(-1, -5) == -1
>>> assert greater(7, 7) == 7
>>> assert greater(0, -10) == 0
```


Tests for q8 (points: None)
```
>>> assert check_square(3, 9) == True
>>> assert check_square(4, 16) == True
>>> assert check_square(5, 25) == True
>>> assert check_square(2, 5) == False
>>> assert check_square(-3, 9) == True
>>> assert check_square(0, 0) == True
>>> assert check_square(1, 1) == True
```


Tests for q9 (points: None)
```
>>> assert ten_tenths(10) == (1.0, 100)
>>> assert ten_tenths(0) == (0.0, 0)
>>> assert ten_tenths(-5) == (-0.5, -50)
>>> assert ten_tenths(100) == (10.0, 1000)
>>> assert ten_tenths(1.5) == (0.15, 15.0)
```
